<img src="logo.png" height="200" width="900"> 

#  Сбор данных: грязная работа вашими руками 

Пришло время самостоятельно написать парсер! Мы будем собирать данные [о ценах на книги.](http://books.toscrape.com)

<!-- grade_id: cell-a1d057ec06e84bdd, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
    
# Подгрузите все необходимые для работы пакеты.
import numpy as np
import pandas as pd
import requests
import bs4

# Если ваш код будет ругаться, что нет пакета lxml, установите его 
# Для этого выполните в одной из ячеек команду !pip3 install lxml
from typing import List

# your code here

```

Прогуляйтесь на сайт http://books.toscrape.com/ и изучите его структуру.  


# 1. Сбор ссылок на книги

Напишите функцию `get_soup`, которая по ссылке возвращает html-разметку страницы в формате `bs4` 

<!-- grade_id: cell-bb3dd3de95e8b98e, points: 0, grade: false, solution: true -->
```python
def get_page_soup(url_link):
    responce = requests.get(url_link)
    if (responce.ok):
        html = responce.content
        soup = bs4.BeautifulSoup(html,'html.parser')
    else:
        print(responce)
        soup = None
    return soup
```

```python
main_url = 'http://books.toscrape.com/catalogue/'
page_number = 5

soup = get_page_soup(main_url + f'page-{page_number}.html')
soup.find_all('a')[0].text
```

Output:
```
'Books to Scrape'
```

```python
books = soup.findAll(lambda tag: tag.name == 'article' and tag.get('class') == ["product_pod"])
[link.h3.a.attrs['href'] for link in books]
```

Output:
```
['princess-jellyfish-2-in-1-omnibus-vol-01-princess-jellyfish-2-in-1-omnibus-1_920/index.html',
 'princess-between-worlds-wide-awake-princess-5_919/index.html',
 'pop-gun-war-volume-1-gift_918/index.html',
 'political-suicide-missteps-peccadilloes-bad-calls-backroom-hijinx-sordid-pasts-rotten-breaks-and-just-plain-dumb-mistakes-in-the-annals-of-american-politics_917/index.html',
 'patience_916/index.html',
 'outcast-vol-1-a-darkness-surrounds-him-outcast-1_915/index.html',
 'orange-the-complete-collection-1-orange-the-complete-collection-1_914/index.html',
 'online-marketing-for-busy-authors-a-step-by-step-guide_913/index.html',
 'on-a-midnight-clear_912/index.html',
 'obsidian-lux-1_911/index.html',
 'my-paris-kitchen-recipes-and-stories_910/index.html',
 'masks-and-shadows_909/index.html',
 'mama-tried-traditional-italian-cooking-for-the-screwed-crude-vegan-and-tattooed_908/index.html',
 'lumberjanes-vol-2-friendship-to-the-max-lumberjanes-5-8_907/index.html',
 'lumberjanes-vol-1-bew
```

Напишите функцию `get_books_links`, которая находит в html-разметке страницы ссылки на странички с отдельными книгами. 

<!-- grade_id: cell-565bf344d89fef43, points: 0, grade: false, solution: true -->
```python
def get_book_links(page_soup) -> List[str]:
    books = soup.findAll(lambda tag: tag.name == 'article' and tag.get('class') == ["product_pod"])
    return [link.h3.a.attrs['href'] for link in books]
```

С помощью цикла соберите в лист `book_links` первые 200 книг.

<!-- grade_id: cell-883211fab847e672, points: 0, grade: false, solution: true -->
```python
book_links = []
page_number = 1
while len(book_links) < 200:
    link = main_url + 'page-{}.html'.format(page_number)
    soup = get_page_soup(link)
    res = get_book_links(soup)
    book_links.extend(res)
    page_number += 1
```

```python
book_links = sorted(book_links)
```

<!-- grade_id: cell-fb6eac2f95e1147e, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert len(book_links) == 200
```

# 2. Сбор информации о книгах 

Напишите несколько небольших функций, которые собирают различные данные об одной книге, необходимые для ответов на вопросы ниже. Информацию о книге собирайте в виде словаря вида 

```
{ 'name': 'Преступление и наказание', 'rating': 1, 'description': 'ужасно депрессивная книга', ... }

```

<!-- grade_id: cell-a3b3b4c91addb471, points: 0, grade: false, solution: true -->
```python
def get_name(soup):
    describe = soup.find('div', attrs={'class':'col-sm-6 product_main'})
    result = describe.h1.text
    return result

def get_rating(soup):
    describe = soup.find('p', attrs={'class':'star-rating'})
    answer = describe.get('class')[1]
    if (answer == 'One'):
        return 1
    elif (answer == 'Two'):
        return 2
    elif (answer == 'Three'):
        return 3
    elif (answer == 'Four'):
        return 4
    elif (answer == 'Five'):
        return 5
    else:
        return 0
```

```python
def get_descr(soup):
    describe = soup.findAll('p')[3]
    return describe.text
```

```python
def get_tax(soup):
    describe = soup.findAll('td')[4]
    return float(describe.text[1:])
```

```python
def get_price(soup):
    describe = soup.findAll('td')[2]
    return float(describe.text[1:])
```

```python
def get_attr(link):
    result = {}
    main_link = main_url + link
    soup = get_page_soup(main_link)
    result['name'] = get_name(soup)
    result['rating'] = get_rating(soup)
    result['description'] = get_descr(soup).replace("\n", '')
    result['tax'] = get_tax(soup)
    result['price'] = get_price(soup)
    return result
```

```python
get_attr("the-bridge-to-consciousness-im-writing-the-bridge-between-science-and-our-old-and-new-beliefs_840/index.html")
```

Output:
```
{'name': "The Bridge to Consciousness: I'm Writing the Bridge Between Science and Our Old and New Beliefs.",
 'rating': 3,
 'description': '',
 'tax': 0.0,
 'price': 32.0}
```

Пройдите циклом по всем сыслкам из списка `book_links` и соберите данные о книгах в вектор `book_info`. 

<!-- grade_id: cell-240b7f956c433b45, points: 0, grade: false, solution: true -->
```python
book_info = []
for link in book_links:
    book_info.append(get_attr(link))
```

Превратим вектор из информации в полноценную таблицу с данными. 

```python
book_info_df = pd.DataFrame(book_info)
print(book_info_df.shape)
book_info_df.head()
```

Output:
```
(200, 5)
                                                name  rating  \
0  A Court of Thorns and Roses (A Court of Thorns...       1   
1                         A Fierce and Subtle Poison       4   
2            A Flight of Arrows (The Pathfinders #2)       5   
3                               A Light in the Attic       3   
4                                   A Murder in Time       1   

                                         description  tax  price  
0  A thrilling, seductive new series from New Yor...  0.0  52.37  
1  Everyone knows the legends about the cursed gi...  0.0  28.13  
2  October 1776--August 1777It is said that what ...  0.0  55.53  
3  It's hard to imagine a world without A Light i...  0.0  51.77  
4  Beautiful and brilliant, Kendra Donovan is a r...  0.0  16.64
```

```python
book_info_df
```

Output:
```
name  rating  \
0    A Court of Thorns and Roses (A Court of Thorns...       1   
1                           A Fierce and Subtle Poison       4   
2              A Flight of Arrows (The Pathfinders #2)       5   
3                                 A Light in the Attic       3   
4                                     A Murder in Time       1   
..                                                 ...     ...   
195  William Shakespeare's Star Wars: Verily, A New...       4   
196                    Without Borders (Wanderlove #1)       2   
197  Worlds Elsewhere: Journeys Around Shakespeare’...       5   
198  You Are What You Love: The Spiritual Power of ...       4   
199                     You can't bury them all: Poems       2   

                                           description  tax  price  
0    A thrilling, seductive new series from New Yor...  0.0  52.37  
1    Everyone knows the legends about the cursed gi...  0.0  28.13  
2    October 1776--August 1777It is said that what
```

Теперь, когда все данные собраны, настало время ответить на несколько вопросов:

- У скольких книг отсутствует описание? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans1`. 

<!-- grade_id: cell-818d91b73f034e1a, points: 0, grade: false, solution: true -->
```python
ans1 = book_info_df[book_info_df["description"] == ''].shape[0]
```

```python
print(ans1)
```

Output:
```
1
```

<!-- grade_id: cell-4b9e7e39b8921ba5, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
```

- Сколько раз в данных встречается налог, больший нуля?  Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans2`. 

<!-- grade_id: cell-be342205fa946a79, points: 0, grade: false, solution: true -->
```python
ans2 = book_info_df[book_info_df["tax"] > 0].shape[0]
```

```python
print(ans2)
```

Output:
```
0
```

<!-- grade_id: cell-c3b1e5331d25a061, points: 0, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
```

- Сколько раз рейтинг книги составлял пять звезд? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans3`. 

<!-- grade_id: cell-a904ff9c737f3bae, points: 0, grade: false, solution: true -->
```python
ans3 = book_info_df[book_info_df["rating"] == 5].shape[0]
```

```python
print(ans3)
```

Output:
```
37
```

<!-- grade_id: cell-ac096895d662ffc8, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно

```

- Какова средняя цена книг (без учета налога)? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans4`. 

<!-- grade_id: cell-60c854eb83c66df8, points: 0, grade: false, solution: true -->
```python
ans4 = book_info_df["price"].mean()
```

```python
print(ans4)
```

Output:
```
34.79625000000001
```

<!-- grade_id: cell-36295b0dc83ec0ad, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно

```